# `evaluation_interval` is not a valid top-level key.
# `rule_files` is correct.
rule_files:
  - 'prometheus.api_server_alerts.yaml'

tests:
  # Test Increasing4xx5xxErrorRate alert - FIRING
  - name: Test Increasing4xx5xxErrorRate Firing
    interval: 1m
    input_series:
      # Current 5m error rate is elevated (0.15 = 15%)
      - series: 'api_server_error_rate_5m{service="api-server", source_cluster="kflux-prd-es01"}'
        values: '0.15+0x20'
      # 30-day average is 0.10 (10%)
      - series: 'api_server_error_rate_30d_avg{service="api-server", source_cluster="kflux-prd-es01"}'
        values: '0.10+0x20'
      # 30-day stddev is 0.01 (1%) (so threshold is 0.10 + 2*0.01 = 0.12, and 0.15 > 0.12)
      - series: 'api_server_error_rate_30d_stddev{service="api-server", source_cluster="kflux-prd-es01"}'
        values: '0.01+0x20'

    alert_rule_test:
      - eval_time: 10m
        alertname: Increasing4xx5xxErrorRate
        exp_alerts:
          - exp_labels:
              severity: warning
              slo: "false"
              service: "api-server"
              source_cluster: "kflux-prd-es01"
            exp_annotations:
              summary: "API Server 4xx/5xx error rate in cluster kflux-prd-es01"
              description: "API Server is experiencing more than 10% failures based on the 30d average and 2x standard deviation in cluster kflux-prd-es01"
              alert_routing_key: <@U03TYEVLM71>
              team: o11y
              runbook_url: null

  # Test Increasing4xx5xxErrorRate alert - NOT FIRING
  - name: Test Increasing4xx5xxErrorRate Not Firing
    interval: 1m
    input_series:
      # Current 5m error rate is 0.11 (11%, within threshold)
      - series: 'api_server_error_rate_5m{service="api-server", source_cluster="kflux-prd-es01"}'
        values: '0.11+0x20'
      # 30-day average is 0.10 (10%)
      - series: 'api_server_error_rate_30d_avg{service="api-server", source_cluster="kflux-prd-es01"}'
        values: '0.10+0x20'
      # 30-day stddev is 0.01 (1%) (so threshold is 0.10 + 2*0.01 = 0.12, and 0.11 < 0.12)
      - series: 'api_server_error_rate_30d_stddev{service="api-server", source_cluster="kflux-prd-es01"}'
        values: '0.01+0x20'

    alert_rule_test:
      - eval_time: 10m
        alertname: Increasing4xx5xxErrorRate
        exp_alerts: []

  # Start of a new test case. The interval and input series must be under this item.
  - name: Test KonfluxAPIServerHighErrorRate
    interval: 1m # Correct placement and spelling of interval
    input_series:
      - series: 'code:apiserver_request_total:rate5m{code="503", source_cluster="kflux-prd-es01"}'
        values: '0+10x60'
      - series: 'code:apiserver_request_total:rate5m{code="429", source_cluster="kflux-prd-es01"}'
        values: '0+5x60'
      - series: 'code:apiserver_request_total:rate5m{code="200", source_cluster="kflux-prd-es01"}'
        values: '0+50x60'

    # This block tests the alert's state (PENDING/FIRING).
    alert_rule_test:
      # This checks for a FIRING alert. It must be checked AFTER the 'for' duration.
      - eval_time: 16m
        alertname: KonfluxAPIServerHighErrorRate
        exp_alerts:
          - exp_labels:
              severity: critical
              slo: "true"
              source_cluster: "kflux-prd-es01"
            exp_annotations:
              summary: "API Server 5XX error count in cluster kflux-prd-es01"
              description: "API Server is experiencing 10% of 5XX failures based on the total requests for 15min in cluster kflux-prd-es01"
              alert_team_handle: "<!subteam^S05Q1P4Q2TG>"
              team: konflux-infra
              runbook_url: https://gitlab.cee.redhat.com/konflux/docs/sop/-/blob/main/infra/sre/api_server_error_rate_alert.md


  # High CPU usage - firing
  #- name: Test KonfluxAPIServerHighCPUUsage
  #  interval: 1m
  #  input_series:
  #    - series: 'process_cpu_seconds_total{job="apiserver",source_cluster="cluster01"}'
  #      # Value that produces a rate of 500/min.
  #      values: '0 500 1000 1500 2000 2500 3000 3500 4000 4500'
  #    - series: 'node_cpu_seconds_total{mode="idle",source_cluster="cluster01"}'
  #      # Value that produces a rate of 550/min. Ratio is ~91%, which is > 80%.
  #      values: '0 550 1100 1650 2200 2750 3300 3850 4400 4950'

  # alert_rule_test:
  #    - alertname: KonfluxAPIServerHighCPUUsage
  #      eval_time: 10m
  #      exp_alerts:
  #        - exp_labels:
  #            severity: critical
  #            slo: "true"
  #            source_cluster: "cluster01"
  #          exp_annotations:
  #            summary: API Server CPU usage is above 80% in cluster01
  #            description: >-
  #              API Server process CPU consumption is above 80% for 5 minutes in cluster cluster01.
  #            alert_team_handle: <!subteam^S05Q1P4Q2TG>
  #            team: konflux-infra
  #            runbook_url: https://gitlab.cee.redhat.com/konflux/docs/sop/-/blob/main/infra/sre/apiserver_HighCPUUsage.md


  ## High CPU usage - NOT firing
  #- name: Test KonfluxAPIServerHighCPUUsage
  #  interval: 1m
  #  input_series:
  #    - series: 'process_cpu_seconds_total{job="apiserver",source_cluster="cluster01"}'
  #      values: '0 0 0 0 0 0 0 0 0 0'
  #    - series: 'node_cpu_seconds_total{mode="idle",source_cluster="cluster01"}'
  #      values: '0 1 2 3 4 5 6 7 8 9'
  #  alert_rule_test:
  #    - alertname: KonfluxAPIServerHighCPUUsage
  #      eval_time: 10m
  #      exp_alerts: []

# High memory usage - firing
  - name: Test KonfluxAPIServerHighMemoryUsageFiring
    interval: 1m
    input_series:
      - series: 'process_resident_memory_bytes{job="apiserver",source_cluster="cluster01"}'
        values: '60x60'
      - series: 'node_memory_MemTotal_bytes{source_cluster="cluster01"}'
        values: '70x60'
    alert_rule_test:
      - alertname: KonfluxAPIServerHighMemoryUsage
        eval_time: 10m
        exp_alerts:
          - exp_labels:
              severity: critical
              slo: "true"
              source_cluster: "cluster01"
            exp_annotations:
              summary: API Server memory usage is above 80% in cluster01
              description: >-
                API Server process memory consumption is above 80% for 5 minutes in cluster cluster01.
              alert_team_handle: <!subteam^S05Q1P4Q2TG>
              team: konflux-infra
              runbook_url: https://gitlab.cee.redhat.com/konflux/docs/sop/-/blob/main/infra/sre/apiserver_HighMemoryUsage.md

  # High memory usage - NOT firing
  - name: Test KonfluxAPIServerHighMemoryUsage
    interval: 1m
    input_series:
      - series: 'process_resident_memory_bytes{job="apiserver",source_cluster="cluster01"}'
        values: '6e10x10'
      - series: 'node_memory_MemTotal_bytes{source_cluster="cluster01"}'
        values: '1e11x10'
    alert_rule_test:
      - alertname: KonfluxAPIServerHighMemoryUsage
        eval_time: 10m
        exp_alerts: []