evaluation_interval: 1m

rule_files:
  - prometheus.application_alerts.yaml

tests:
  # ----- Application Deletion Alerting Tests -----
  - interval: 1m
    input_series:

      # HAS is experiencing application deletion failure rate of 100%, so it will be alerted
      - series: 'has_application_deletion_total{namespace="application-service", pod="has", source_cluster="cluster01"}'
        values: '1+1x64'
      - series: 'has_application_failed_deletion_total{namespace="application-service", pod="has", source_cluster="cluster01"}'
        values: '1+1x64'

    alert_rule_test:
      - eval_time: 65m
        alertname: ApplicationDeletionErrors
        exp_alerts:
          - exp_labels:
              severity: critical
              slo: "true"
              namespace: application-service
              pod: has
              source_cluster: cluster01
            exp_annotations:
              summary: >-
                HAS is experiencing application deletion failures of >5%
              description: >-
                Application controller in Pod has for namespace
                application-service on cluster cluster01 is failing to
                successfully delete at least 95% of applications over the past hour
              alert_team_handle: <!subteam^S04MSCVRF4Z>
              team: has
              runbook_url: https://gitlab.cee.redhat.com/konflux/docs/sop/-/blob/main/has/application-delete-failed.md

  - interval: 1m
    input_series:

      # HAS experienced ~10% (10.9%) deletion failure rate, should still alert
      - series: 'has_application_deletion_total{namespace="application-service", pod="has", source_cluster="cluster01"}'
        values: '10+10x64'
      - series: 'has_application_failed_deletion_total{namespace="application-service", pod="has", source_cluster="cluster01"}'
        values: '0x20 0+10x7 70x38'

    alert_rule_test:
      - eval_time: 65m
        alertname: ApplicationDeletionErrors
        exp_alerts:
          - exp_labels:
              severity: critical
              slo: "true"
              namespace: application-service
              pod: has
              source_cluster: cluster01
            exp_annotations:
              summary: >-
                HAS is experiencing application deletion failures of >5%
              description: >-
                Application controller in Pod has for namespace
                application-service on cluster cluster01 is failing to
                successfully delete at least 95% of applications over the past hour
              alert_team_handle: <!subteam^S04MSCVRF4Z>
              team: has
              runbook_url: https://gitlab.cee.redhat.com/konflux/docs/sop/-/blob/main/has/application-delete-failed.md

  - interval: 1m
    input_series:

      # HAS experienced 5% deletion failure rate, should not alert
      - series: 'has_application_deletion_total{namespace="application-service", pod="has", source_cluster="cluster01"}'
        values: '20+20x64'
      - series: 'has_application_failed_deletion_total{namespace="application-service", pod="has", source_cluster="cluster01"}'
        values: '1+1x64'

    alert_rule_test:
      - eval_time: 65m
        alertname: ApplicationDeletionErrors

  - interval: 1m
    input_series:

      # HAS experienced ~1% deletion failure rate, should not alert
      - series: 'has_application_deletion_total{namespace="application-service", pod="has", source_cluster="cluster01"}'
        values: '0+10x64'
      - series: 'has_application_failed_deletion_total{namespace="application-service", pod="has", source_cluster="cluster01"}'
        values: '0x9 5x54'

    alert_rule_test:
      - eval_time: 65m
        alertname: ApplicationDeletionErrors


  # ----- Application Creation Alerting Tests -----
  - interval: 1m
    input_series:

      # HAS is experiencing application creation failure rate of 100%, so it will be alerted
      - series: 'has_application_creation_total{namespace="application-service", pod="has", source_cluster="cluster01"}'
        values: '1+1x64'
      - series: 'has_application_failed_creation_total{namespace="application-service", pod="has", source_cluster="cluster01"}'
        values: '1+1x64'

    alert_rule_test:
      - eval_time: 65m
        alertname: ApplicationCreationErrors
        exp_alerts:
          - exp_labels:
              severity: critical
              slo: "true"
              namespace: application-service
              pod: has
              source_cluster: cluster01
            exp_annotations:
              summary: >-
                HAS is experiencing application creation failures of >5%
              description: >-
                Application controller in Pod has for namespace
                application-service on cluster cluster01 is failing to
                successfully create at least 95% of applications over the past hour
              alert_team_handle: <!subteam^S04MSCVRF4Z>
              team: has
              runbook_url: https://gitlab.cee.redhat.com/konflux/docs/sop/-/blob/main/has/application-create-failed.md

  - interval: 1m
    input_series:

      # HAS experienced ~10% (10.9%) creation failure rate, should still alert
      - series: 'has_application_creation_total{namespace="application-service", pod="has", source_cluster="cluster01"}'
        values: '10+10x64'
      - series: 'has_application_failed_creation_total{namespace="application-service", pod="has", source_cluster="cluster01"}'
        values: '0x20 0+10x7 70x38'

    alert_rule_test:
      - eval_time: 65m
        alertname: ApplicationCreationErrors
        exp_alerts:
          - exp_labels:
              severity: critical
              slo: "true"
              namespace: application-service
              pod: has
              source_cluster: cluster01
            exp_annotations:
              summary: >-
                HAS is experiencing application creation failures of >5%
              description: >-
                Application controller in Pod has for namespace
                application-service on cluster cluster01 is failing to
                successfully create at least 95% of applications over the past hour
              alert_team_handle: <!subteam^S04MSCVRF4Z>
              team: has
              runbook_url: https://gitlab.cee.redhat.com/konflux/docs/sop/-/blob/main/has/application-create-failed.md

  - interval: 1m
    input_series:

      # HAS experienced 5% creation failure rate, should not alert
      - series: 'has_application_creation_total{namespace="application-service", pod="has", source_cluster="cluster01"}'
        values: '20+20x64'
      - series: 'has_application_failed_creation_total{namespace="application-service", pod="has", source_cluster="cluster01"}'
        values: '1+1x64'

    alert_rule_test:
      - eval_time: 65m
        alertname: ApplicationCreationErrors

  - interval: 1m
    input_series:

      # HAS experienced ~1% creation failure rate, should not alert
      - series: 'has_application_creation_total{namespace="application-service", pod="has", source_cluster="cluster01"}'
        values: '0+10x64'
      - series: 'has_application_failed_creation_total{namespace="application-service", pod="has", source_cluster="cluster01"}'
        values: '0x9 5x54'

    alert_rule_test:
      - eval_time: 65m
        alertname: ApplicationCreationErrors
