evaluation_interval: 1m

rule_files:
  - 'prometheus.high_memory_usage_alert.yaml'

tests:
  - interval: 1m
    input_series:
      - series: "kube_pod_container_memory_usage_bytes{tenant_id='tenant', pod='pod', namespace='namespace', container='container', source_cluster='cluster01'}"
        values: '1x60'
      - series: "kube_pod_container_resource_limits_memory_bytes{tenant_id='tenant', pod='pod', namespace='namespace', container='container', source_cluster='cluster01'}"
        values: "1x60"

    alert_rule_test:
      - eval_time: 5m
        alertname: HighMemoryUsageAlert
        exp_alerts:
          - exp_labels:
              namespace: "namespace"
              pod: "pod"
              severity: "warning"
              tenant_id: "tenant"
              container: "container"
              source_cluster: "cluster01"
            exp_annotations:
              description: "Pod pod for namespace namespace on cluster cluster01 has high memory usage"
              alert_routing_key: "namespace"
              summary: "Pod pod has high memory usage"
              runbook_url: 'https://gitlab.cee.redhat.com/konflux/docs/sop/-/blob/main/sre/alert-rule-HighMemoryAlertUsage.md'
