evaluation_interval: 1m

rule_files:
  - prometheus.kyverno_alerts.yaml

tests:
# --- alerts for kyverno pods failures ---

  # Scenario 1: Deployment is continuously down for 3 minutes, alert should fire.
  - interval: 1m
    input_series:
      - series: 'kube_deployment_spec_replicas{deployment="kyverno-admission-controller", namespace="konflux-kyverno", source_cluster="stone-stg-rh01"}'
        values: '1 1 1 1 1 1'
      - series: 'kube_deployment_status_replicas_ready{deployment="kyverno-admission-controller", namespace="konflux-kyverno", source_cluster="stone-stg-rh01"}'
        values: '0 0 0 0 0 0'
      - series: 'kube_deployment_spec_replicas{deployment="kyverno-background-controller", namespace="konflux-kyverno", source_cluster="stone-stg-rh01"}'
        values: '1 1 1 1 1 1'
      - series: 'kube_deployment_status_replicas_ready{deployment="kyverno-background-controller", namespace="konflux-kyverno", source_cluster="stone-stg-rh01"}'
        values: '1 1 1 1 1 1'
    alert_rule_test:
      - eval_time: 2m
        alertname: KyvernoDeploymentDown
        exp_alerts: []
      - eval_time: 3m
        alertname: KyvernoDeploymentDown
        exp_alerts:
          - exp_labels:
              severity: critical
              slo: "true"
              deployment: kyverno-admission-controller
              namespace: konflux-kyverno
              source_cluster: stone-stg-rh01
            exp_annotations:
              summary: Kyverno has reported that not all the deployments pods are up and running for the last 3 minutes.
              description: "The Kyverno Deployment 'kyverno-admission-controller' in namespace 'konflux-kyverno' on cluster 'stone-stg-rh01' is unhealthy."
              team: konflux-infra
              alert_team_handle: <!subteam^S05Q1P4Q2TG>
              runbook_url: https://gitlab.cee.redhat.com/konflux/docs/sop/-/blob/main/kyverno/README.md
      - eval_time: 4m
        alertname: KyvernoDeploymentDown
        exp_alerts:
          - exp_labels:
              severity: critical
              slo: "true"
              deployment: kyverno-admission-controller
              namespace: konflux-kyverno
              source_cluster: stone-stg-rh01
            exp_annotations:
              summary: Kyverno has reported that not all the deployments pods are up and running for the last 3 minutes.
              description: "The Kyverno Deployment 'kyverno-admission-controller' in namespace 'konflux-kyverno' on cluster 'stone-stg-rh01' is unhealthy."
              team: konflux-infra
              alert_team_handle: <!subteam^S05Q1P4Q2TG>
              runbook_url: https://gitlab.cee.redhat.com/konflux/docs/sop/-/blob/main/kyverno/README.md
  # Scenario 2: Deployment is down for less than 3 minutes, alert should NOT fire.
  - interval: 1m
    input_series:
      - series: 'kube_deployment_spec_replicas{deployment="kyverno-background-controller", namespace="konflux-kyverno", source_cluster="stone-stg-rh01"}'
        values: '1 1 1 1 1'
      - series: 'kube_deployment_status_replicas_ready{deployment="kyverno-background-controller", namespace="konflux-kyverno", source_cluster="stone-stg-rh01"}'
        values: '0 0 1 0 1'
    alert_rule_test:
      - eval_time: 2m
        alertname: KyvernoDeploymentDown
        exp_alerts: []
      - eval_time: 4m
        alertname: KyvernoDeploymentDown
        exp_alerts: []

# --- alerts for kyverno controller requeue errors ---

  # Scenario 1: Kyverno controller is continuously experiencing requeue errors for 3 minutes, alert should fire.
  - interval: 1m
    input_series:
      - series: 'kyverno_controller_requeue_total{controller_name="certmanager-controller", source_cluster="stone-stg-rh01"}'
        values: '0 1 2 3 4 4'
    alert_rule_test:
      - eval_time: 2m
        alertname: KyvernoControllerRequeueErrors
        exp_alerts: []
      - eval_time: 4m
        alertname: KyvernoControllerRequeueErrors
        exp_alerts:
          - exp_labels:
              severity: critical
              slo: "true"
              controller_name: certmanager-controller
              source_cluster: stone-stg-rh01
            exp_annotations:
              summary: Kyverno internal controllers are encountering errors causing items to be re-queued.
              description: "The Kyverno controller 'certmanager-controller' on cluster 'stone-stg-rh01' is experiencing a consistent rate of requeue errors."
              team: konflux-infra
              alert_team_handle: <!subteam^S05Q1P4Q2TG>
              runbook_url: https://gitlab.cee.redhat.com/konflux/docs/sop/-/blob/main/kyverno/README.md
      - eval_time: 5m
        alertname: KyvernoControllerRequeueErrors
        exp_alerts: []
  # Scenario 2: Kyverno controller is experiencing requeue errors for less than 3 minutes, alert should NOT fire.
  - interval: 1m
    input_series:
      - series: 'kyverno_controller_requeue_total{controller_name="certmanager-controller", source_cluster="stone-stg-rh01"}'
        values: '0 1 0 0 0 0'
    alert_rule_test:
      - eval_time: 2m
        alertname: KyvernoControllerRequeueErrors
        exp_alerts: []
      - eval_time: 4m
        alertname: KyvernoControllerRequeueErrors
        exp_alerts: []

# --- alerts for kyverno policy failures ---

  # Scenario 1: Kyverno policies are failing or erroring out more often than usual for 3 minutes, alert should fire.
  - interval: 1m
    input_series:
      - series: 'kyverno_policy_results_total{policy_name="propagate-cost-management-labels",  rule_name="autogen-propagate-existing-cost-center-from-namespace", rule_result="error", instance="10.128.10.157:8000", source_cluster="stone-stg-rh01"}'
        values: '0 10 35 70 105 120 120 135 160 180 191'
    alert_rule_test:
      - eval_time: 2m
        alertname: KyvernoPolicyFailures
        exp_alerts: []
      - eval_time: 4m
        alertname: KyvernoPolicyFailures
        exp_alerts: []
      - eval_time: 5m
        alertname: KyvernoPolicyFailures
        exp_alerts:
          - exp_labels:
              severity: warning
              slo: "false"
              policy_name: propagate-cost-management-labels
              rule_name: autogen-propagate-existing-cost-center-from-namespace
              instance: "10.128.10.157:8000"
              source_cluster: stone-stg-rh01
              rule_result: error
            exp_annotations:
              summary: Kyverno policies are failing or erroring out more often than usual.
              description: "Kyverno policy 'propagate-cost-management-labels' with rule 'autogen-propagate-existing-cost-center-from-namespace' is encountering internal errors during evaluation or rejecting resources on instance '10.128.10.157:8000' on cluster 'stone-stg-rh01'."
              team: konflux-infra
              alert_routing_key: infra
              runbook_url: https://gitlab.cee.redhat.com/konflux/docs/sop/-/blob/main/kyverno/README.md
      - eval_time: 6m
        alertname: KyvernoPolicyFailures
        exp_alerts: []
      - eval_time: 9m
        alertname: KyvernoPolicyFailures
        exp_alerts: []
      - eval_time: 10m
        alertname: KyvernoPolicyFailures
        exp_alerts:
          - exp_labels:
              severity: warning
              slo: "false"
              policy_name: propagate-cost-management-labels
              rule_name: autogen-propagate-existing-cost-center-from-namespace
              instance: "10.128.10.157:8000"
              source_cluster: stone-stg-rh01
              rule_result: error
            exp_annotations:
              summary: Kyverno policies are failing or erroring out more often than usual.
              description: "Kyverno policy 'propagate-cost-management-labels' with rule 'autogen-propagate-existing-cost-center-from-namespace' is encountering internal errors during evaluation or rejecting resources on instance '10.128.10.157:8000' on cluster 'stone-stg-rh01'."
              team: konflux-infra
              alert_routing_key: infra
              runbook_url: https://gitlab.cee.redhat.com/konflux/docs/sop/-/blob/main/kyverno/README.md
  # Scenario 2: Kyverno policies are failing or erroring out for less than 3 minutes, alert should NOT fire.
  - interval: 1m
    input_series:
      - series: 'kyverno_policy_results_total{policy_name="propagate-cost-management-labels", policy_rule="autogen-propagate-existing-cost-center-from-namespace", rule_result="error", instance="10.128.10.157:8000", source_cluster="stone-stg-rh01"}'
        values: '10 20 30 40 50 60'
    alert_rule_test:
      - eval_time: 2m
        alertname: KyvernoPolicyFailures
        exp_alerts: []
      - eval_time: 4m
        alertname: KyvernoPolicyFailures
        exp_alerts: []