evaluation_interval: 1m

rule_files:
  - 'prometheus.mintmaker_controller_up_alerts.yaml'

tests:
  - interval: 1m
    input_series:
      - series: 'konflux_up{namespace="mintmaker", check="replicas-available", service="mintmaker-controller-manager", source_cluster="c1"}'
        values: '0 0 0 0 0 0 0 0 0 0 0 0 0 0 0'
      - series: 'konflux_up{namespace="mintmaker", check="replicas-available", service="mintmaker-controller-manager", source_cluster="c2"}'
        values: '1 1 1 1 1 1 1 1 1 1 1 1 1 1 1'

    alert_rule_test:
      - eval_time: 15m
        alertname: MintMakerControllerDown
        exp_alerts:
          - exp_labels:
              severity: critical
              check: replicas-available
              namespace: mintmaker
              service: mintmaker-controller-manager
              slo: "true"
              source_cluster: c1
            exp_annotations:
              summary: MintMaker controller is down in cluster c1
              description: >
                The MintMaker controller pod has been down for 5min in cluster c1
              alert_team_handle: <!subteam^S078T8TMQP5>
              team: mintmaker
              runbook_url: https://gitlab.cee.redhat.com/konflux/docs/sop/-/blob/main/mintmaker/faq.md

  - interval: 1m
    input_series:
      - series: 'konflux_up{namespace="mintmaker", check="replicas-available", service="redis", source_cluster="c5"}'
        values: '0 0 0 0 0 0 0 0 0 0 0 0 0 0 0'
      - series: 'konflux_up{namespace="mintmaker", check="replicas-available", service="redis", source_cluster="c6"}'
        values: '1 1 1 1 1 1 1 1 1 1 1 1 1 1 1'

    alert_rule_test:
      - eval_time: 15m
        alertname: MintMakerCacheDown
        exp_alerts:
          - exp_labels:
              severity: critical
              check: replicas-available
              namespace: mintmaker
              service: redis
              slo: "true"
              source_cluster: c5
            exp_annotations:
              summary: MintMaker cache is down in cluster c5
              description: >
                The MintMaker cache pod has been down for 5min in cluster c5
              alert_team_handle: <!subteam^S078T8TMQP5>
              team: mintmaker
              runbook_url: https://gitlab.cee.redhat.com/konflux/docs/sop/-/blob/main/mintmaker/faq.md

  - interval: 1m
    input_series:
      - series: 'konflux_up{namespace="mintmaker", check="replicas-available", service="mintmaker-controller-manager", source_cluster="stone-prod-p02"}'
        values: '0 0 0 0 0'
      - series: 'konflux_up{namespace="mintmaker", check="replicas-available", service="mintmaker-controller-manager", source_cluster="c7"}'
        values: '1 1 1 1 1'

    alert_rule_test:
      - eval_time: 5m
        alertname: MintMakerControllerDown
        exp_alerts: null

  - interval: 1m
    input_series:
      - series: 'konflux_up{namespace="mintmaker", check="replicas-available", service="redis", source_cluster="stone-prod-p02"}'
        values: '0 0 0 0 0'
      - series: 'konflux_up{namespace="mintmaker", check="replicas-available", service="redis", source_cluster="c8"}'
        values: '1 1 1 1 1'

    alert_rule_test:
      - eval_time: 5m
        alertname: MintMakerControllerDown
        exp_alerts: null

  - interval: 1m
    input_series:
      # DUC created at test start time (timestamp 0), value stays constant
      - series: 'mintmaker_dependency_update_check_creation_time{name="dependencyupdatecheck-old",namespace="mintmaker", source_cluster="c7"}'
        values: '0+0x300'

    alert_rule_test:
      # Evaluate at 4.5 hours - should NOT fire yet (exactly at threshold)
      - eval_time: 4h30m
        alertname: MintMakerDependencyUpdateCheckStale
        exp_alerts: []

      # Evaluate at 4.5 hours + 5 minutes - should fire after 5m delay
      - eval_time: 4h36m
        alertname: MintMakerDependencyUpdateCheckStale
        exp_alerts:
          - exp_labels:
              severity: warning
              slo: "false"
              source_cluster: c7
            exp_annotations:
              summary: "MintMaker dependency update check creation is stale in cluster c7"
              description: >
                No new DependencyUpdateCheck resources have been created in the last 4.5 hours in cluster c7.
                The last DUC was created more than 4.5 hours ago.
              alert_team_handle: <!subteam^S078T8TMQP5>
              team: mintmaker
              runbook_url: https://gitlab.cee.redhat.com/konflux/docs/sop/-/blob/main/mintmaker/faq.md