evaluation_interval: 1m

rule_files:
  - 'prometheus.multi_platform_alerts.yaml'

tests:
  - interval: 1m
    input_series:
      - series: 'konflux_up{namespace="multi-platform-controller", check="replicas-available", service="multi-platform-controller", source_cluster="c1"}'
        values: '0 0 0 0 0'
      - series: 'konflux_up{namespace="multi-platform-controller", check="replicas-available", service="multi-platform-controller", source_cluster="c2"}'
        values: '1 1 1 1 1'

    alert_rule_test:
      - eval_time: 5m
        alertname: MultiPlatformControllerDown
        exp_alerts:
          - exp_labels:
              severity: critical
              check: replicas-available
              namespace: multi-platform-controller
              service: multi-platform-controller
              slo: "true"
              source_cluster: c1
            exp_annotations:
              summary: Multiplatform controller is down in cluster c1
              description: >
                The multi-platform-controller pod has been down for 5min in cluster c1
              alert_routing_key: infra
              team: konflux-infra
              runbook_url: https://gitlab.cee.redhat.com/konflux/docs/sop/-/blob/main/o11y/alert-rules/alert-rule-MultiPlatformControllerDown.md?ref_type=heads

  - interval: 1m
    input_series:
      - series: 'konflux_up{namespace="multi-platform-controller", check="replicas-available", service="multi-platform-otp-server", source_cluster="c3"}'
        values: '0 0 0 0 0'
      - series: 'konflux_up{namespace="multi-platform-controller", check="replicas-available", service="multi-platform-otp-server", source_cluster="c4"}'
        values: '1 1 1 1 1'

    alert_rule_test:
      - eval_time: 5m
        alertname: MultiPlatformOTPDown
        exp_alerts:
          - exp_labels:
              severity: critical
              check: replicas-available
              namespace: multi-platform-controller
              service: multi-platform-otp-server
              slo: "true"
              source_cluster: c3
            exp_annotations:
              summary: Multiplatform OTP is down in cluster c3
              description: >
                The multi-platform otp pod has been down for 5min in cluster c3
              alert_routing_key: infra
              team: konflux-infra
              runbook_url: https://gitlab.cee.redhat.com/konflux/docs/sop/-/blob/main/o11y/alert-rules/alert-rule-MultiPlatformOTPDown.md?ref_type=heads

  - interval: 1m
    input_series:
      - series: 'multi_platform_controller_available{check="s390x", source_cluster="c5"}'
        values: '0 0 0 0 0'
      - series: 'multi_platform_controller_available{check="ppc64le", source_cluster="c5"}'
        values: '0 0 0 0 0'
      - series: 'multi_platform_controller_available{check="amd64", source_cluster="c6"}'
        values: '1 1 1 1 1'

    alert_rule_test:
      - eval_time: 5m
        alertname: MultiPlatformControllerPlatformUnhealthy
        exp_alerts:
          - exp_labels:
              severity: high
              check: s390x
              slo: "false"
              source_cluster: c5
            exp_annotations:
              summary: Multi-platform controller platform s390x health check failed in cluster c5
              description: >
                The multi-platform controller platform health check "s390x" has been failing for 5min in cluster c5
              alert_routing_key: infra
              team: konflux-infra
              runbook_url: https://gitlab.cee.redhat.com/konflux/docs/sop/-/blob/main/o11y/alert-rules/alert-rule-MultiPlatformControllerPlatformUnhealthy.md?ref_type=heads
          - exp_labels:
              severity: high
              check: ppc64le
              slo: "false"
              source_cluster: c5
            exp_annotations:
              summary: Multi-platform controller platform ppc64le health check failed in cluster c5
              description: >
                The multi-platform controller platform health check "ppc64le" has been failing for 5min in cluster c5
              alert_routing_key: infra
              team: konflux-infra
              runbook_url: https://gitlab.cee.redhat.com/konflux/docs/sop/-/blob/main/o11y/alert-rules/alert-rule-MultiPlatformControllerPlatformUnhealthy.md?ref_type=heads
  - interval: 30s
    input_series:
      - series: 'multi_platform_controller_running_tasks{source_cluster="c7"}'
        values: '-1 -1 -1 -1'
      - series: 'multi_platform_controller_waiting_tasks{source_cluster="c7"}'
        values: '5 5 5 5'
      - series: 'multi_platform_controller_running_tasks{source_cluster="c8"}'
        values: '3 3 3 3'
      - series: 'multi_platform_controller_waiting_tasks{source_cluster="c8"}'
        values: '1 -1 -1 -1'
      - series: 'multi_platform_controller_running_tasks{source_cluster="c9"}'
        values: '1 -1 1 -1'
      - series: 'multi_platform_controller_waiting_tasks{source_cluster="c9"}'
        values: '1 -1 0 1'

    alert_rule_test:
      - eval_time: 2m
        alertname: MultiPlatformControllerMetricsUnhealthy
        exp_alerts:
          - exp_labels:
              severity: warning
              source_cluster: c7
            exp_annotations:
              summary: Multi-platform controller metrics are unhealthy in cluster c7
              description: The multi-platform controller Gauge metrics are showing negative values for the last 1 minute in cluster c7
              alert_routing_key: infra
              team: konflux-infra
          - exp_labels:
              severity: warning
              source_cluster: c8
            exp_annotations:
              summary: Multi-platform controller metrics are unhealthy in cluster c8
              description: The multi-platform controller Gauge metrics are showing negative values for the last 1 minute in cluster c8
              alert_routing_key: infra
              team: konflux-infra

  - interval: 1m
    input_series:
      - series: 'multi_platform_controller_provisioning_successes{source_cluster="c9"}'
        values: '0+24x30 24+0x10'
      - series: 'multi_platform_controller_provisioning_failures{source_cluster="c9"}'
        values: '0+6x30 6+0x10'
      - series: 'multi_platform_controller_provisioning_successes{source_cluster="c10"}'
        values: '0+57x30 57+0x10'
      - series: 'multi_platform_controller_provisioning_failures{source_cluster="c10"}'
        values: '0+3x30 3+0x10'
      - series: 'multi_platform_controller_provisioning_successes{source_cluster="c11"}'
        values: '0+0x40'
      - series: 'multi_platform_controller_provisioning_failures{source_cluster="c11"}'
        values: '0+0x40'

    alert_rule_test:
      - eval_time: 35m
        alertname: MultiPlatformControllerProvisioningSuccessRateTooLow
        exp_alerts:
          - exp_labels:
              severity: warning
              source_cluster: c9
            exp_annotations:
              summary: Multi-platform controller provisioning success rate is too low in cluster c9
              description: The multi-platform controller provisioning success rate is less than 90% for the last 5 minutes in cluster c9
              alert_routing_key: infra
              team: konflux-infra
