evaluation_interval: 1m

rule_files:
  - prometheus.registry_exporter_alerts.yaml

tests:
  # Registry exporter down for quay.io registry
  - interval: 1m
    input_series:
      - series: 'konflux_up{service="registry-exporter", check="quay.io", source_cluster="c1"}'
        values: '0 0 0 0 0 0 0 0 0 0 0 0 0 0 0 0'
      - series: 'konflux_up{service="registry-exporter", check="quay.io", source_cluster="c2"}'
        values: '1 1 1 1 1 1 1 1 1 1 1 1 1 1 1 1'

    alert_rule_test:
      - eval_time: 15m
        alertname: RegistryExporterDown
        exp_alerts:
          - exp_labels:
              severity: critical
              check: quay.io
              service: registry-exporter
              slo: "true"
              source_cluster: c1
            exp_annotations:
              summary: >-
                Registry exporter check failed for registry quay.io in cluster c1
              description: >-
                The registry exporter has failed to successfully test registry quay.io in cluster c1.
                All registry test operations (pull, push, metadata, authentication) have been failing for 15 minutes.
              alert_team_handle: <!subteam^S07SW2EEW3D>
              team: o11y
              runbook_url: https://gitlab.cee.redhat.com/konflux/docs/sop/-/blob/main/o11y/alert-rules/alert-rule-RegistryExporter.md

  # Registry exporter down for multiple registries in same cluster
  - interval: 1m
    input_series:
      - series: 'konflux_up{service="registry-exporter", check="quay.io", source_cluster="c1"}'
        values: '0 0 0 0 0 0 0 0 0 0 0 0 0 0 0 0'
      - series: 'konflux_up{service="registry-exporter", check="registry.example.com", source_cluster="c1"}'
        values: '0 0 0 0 0 0 0 0 0 0 0 0 0 0 0 0'
      - series: 'konflux_up{service="registry-exporter", check="quay.io", source_cluster="c2"}'
        values: '1 1 1 1 1 1 1 1 1 1 1 1 1 1 1 1'

    alert_rule_test:
      - eval_time: 15m
        alertname: RegistryExporterDown
        exp_alerts:
          - exp_labels:
              severity: critical
              check: quay.io
              service: registry-exporter
              slo: "true"
              source_cluster: c1
            exp_annotations:
              summary: >-
                Registry exporter check failed for registry quay.io in cluster c1
              description: >-
                The registry exporter has failed to successfully test registry quay.io in cluster c1.
                All registry test operations (pull, push, metadata, authentication) have been failing for 15 minutes.
              alert_team_handle: <!subteam^S07SW2EEW3D>
              team: o11y
              runbook_url: https://gitlab.cee.redhat.com/konflux/docs/sop/-/blob/main/o11y/alert-rules/alert-rule-RegistryExporter.md
          - exp_labels:
              severity: critical
              check: registry.example.com
              service: registry-exporter
              slo: "true"
              source_cluster: c1
            exp_annotations:
              summary: >-
                Registry exporter check failed for registry registry.example.com in cluster c1
              description: >-
                The registry exporter has failed to successfully test registry registry.example.com in cluster c1.
                All registry test operations (pull, push, metadata, authentication) have been failing for 15 minutes.
              alert_team_handle: <!subteam^S07SW2EEW3D>
              team: o11y
              runbook_url: https://gitlab.cee.redhat.com/konflux/docs/sop/-/blob/main/o11y/alert-rules/alert-rule-RegistryExporter.md

  # Registry exporter recovers before alert fires (should not alert)
  - interval: 1m
    input_series:
      - series: 'konflux_up{service="registry-exporter", check="quay.io", source_cluster="c1"}'
        values: '0 0 0 0 0 0 0 0 0 0 0 0 0 0 1 1'
      - series: 'konflux_up{service="registry-exporter", check="quay.io", source_cluster="c2"}'
        values: '1 1 1 1 1 1 1 1 1 1 1 1 1 1 1 1'

    alert_rule_test:
      - eval_time: 15m
        alertname: RegistryExporterDown
        exp_alerts: []

  # Registry exporter down for different registries in different clusters
  - interval: 1m
    input_series:
      - series: 'konflux_up{service="registry-exporter", check="quay.io", source_cluster="c1"}'
        values: '0 0 0 0 0 0 0 0 0 0 0 0 0 0 0 0'
      - series: 'konflux_up{service="registry-exporter", check="registry.example.com", source_cluster="c2"}'
        values: '0 0 0 0 0 0 0 0 0 0 0 0 0 0 0 0'
      - series: 'konflux_up{service="registry-exporter", check="quay.io", source_cluster="c2"}'
        values: '1 1 1 1 1 1 1 1 1 1 1 1 1 1 1 1'

    alert_rule_test:
      - eval_time: 15m
        alertname: RegistryExporterDown
        exp_alerts:
          - exp_labels:
              severity: critical
              check: quay.io
              service: registry-exporter
              slo: "true"
              source_cluster: c1
            exp_annotations:
              summary: >-
                Registry exporter check failed for registry quay.io in cluster c1
              description: >-
                The registry exporter has failed to successfully test registry quay.io in cluster c1.
                All registry test operations (pull, push, metadata, authentication) have been failing for 15 minutes.
              alert_team_handle: <!subteam^S07SW2EEW3D>
              team: o11y
              runbook_url: https://gitlab.cee.redhat.com/konflux/docs/sop/-/blob/main/o11y/alert-rules/alert-rule-RegistryExporter.md
          - exp_labels:
              severity: critical
              check: registry.example.com
              service: registry-exporter
              slo: "true"
              source_cluster: c2
            exp_annotations:
              summary: >-
                Registry exporter check failed for registry registry.example.com in cluster c2
              description: >-
                The registry exporter has failed to successfully test registry registry.example.com in cluster c2.
                All registry test operations (pull, push, metadata, authentication) have been failing for 15 minutes.
              alert_team_handle: <!subteam^S07SW2EEW3D>
              team: o11y
              runbook_url: https://gitlab.cee.redhat.com/konflux/docs/sop/-/blob/main/o11y/alert-rules/alert-rule-RegistryExporter.md

  # Registry exporter node-level alert tests
  # Node failure when cluster is up (should alert)
  - interval: 1m
    input_series:
      - series: 'konflux_up{service="registry-exporter", check="quay.io", source_cluster="c1"}'
        values: '1x31'
      - series: 'registry_exporter_success{tested_registry="quay.io", node="node1", type="pull", source_cluster="c1"}'
        values: '0x31'
      - series: 'registry_exporter_success{tested_registry="quay.io", node="node1", type="push", source_cluster="c1"}'
        values: '0x31'
      - series: 'registry_exporter_success{tested_registry="quay.io", node="node1", type="metadata", source_cluster="c1"}'
        values: '0x31'
      - series: 'registry_exporter_success{tested_registry="quay.io", node="node1", type="authentication", source_cluster="c1"}'
        values: '0x31'
      - series: 'registry_exporter_success{tested_registry="quay.io", node="node2", type="pull", source_cluster="c1"}'
        values: '1x31'

    alert_rule_test:
      - eval_time: 30m
        alertname: RegistryExporterNodeFailure
        exp_alerts:
          - exp_labels:
              severity: warning
              tested_registry: quay.io
              node: node1
              source_cluster: c1
              slo: "false"
            exp_annotations:
              summary: >-
                Registry exporter running tests on node node1 for registry quay.io in cluster c1 is failing.
              description: >-
                Registry exporter running tests on node node1 for registry quay.io in cluster c1 is failing.
                At least one test type (pull, push, metadata, or authentication) has been failing for 30 minutes.
                The cluster-level check for this registry is still passing, indicating this is a node-specific issue.
              alert_routing_key: o11y
              team: o11y

  # Node failure when cluster is down (should NOT alert)
  - interval: 1m
    input_series:
      - series: 'konflux_up{service="registry-exporter", check="quay.io", source_cluster="c1"}'
        values: '0x31'
      - series: 'registry_exporter_success{tested_registry="quay.io", node="node1", type="pull", source_cluster="c1"}'
        values: '0x31'
      - series: 'registry_exporter_success{tested_registry="quay.io", node="node1", type="push", source_cluster="c1"}'
        values: '0x31'

    alert_rule_test:
      - eval_time: 30m
        alertname: RegistryExporterNodeFailure
        exp_alerts: []

  # Multiple nodes failing in same cluster (should alert for each node)
  - interval: 1m
    input_series:
      - series: 'konflux_up{service="registry-exporter", check="quay.io", source_cluster="c1"}'
        values: '1x31'
      - series: 'registry_exporter_success{tested_registry="quay.io", node="node1", type="pull", source_cluster="c1"}'
        values: '0x31'
      - series: 'registry_exporter_success{tested_registry="quay.io", node="node2", type="pull", source_cluster="c1"}'
        values: '0x31'
      - series: 'registry_exporter_success{tested_registry="quay.io", node="node3", type="pull", source_cluster="c1"}'
        values: '1x31'

    alert_rule_test:
      - eval_time: 30m
        alertname: RegistryExporterNodeFailure
        exp_alerts:
          - exp_labels:
              severity: warning
              tested_registry: quay.io
              node: node1
              source_cluster: c1
              slo: "false"
            exp_annotations:
              summary: >-
                Registry exporter running tests on node node1 for registry quay.io in cluster c1 is failing.
              description: >-
                Registry exporter running tests on node node1 for registry quay.io in cluster c1 is failing.
                At least one test type (pull, push, metadata, or authentication) has been failing for 30 minutes.
                The cluster-level check for this registry is still passing, indicating this is a node-specific issue.
              alert_routing_key: o11y
              team: o11y
          - exp_labels:
              severity: warning
              tested_registry: quay.io
              node: node2
              source_cluster: c1
              slo: "false"
            exp_annotations:
              summary: >-
                Registry exporter running tests on node node2 for registry quay.io in cluster c1 is failing.
              description: >-
                Registry exporter running tests on node node2 for registry quay.io in cluster c1 is failing.
                At least one test type (pull, push, metadata, or authentication) has been failing for 30 minutes.
                The cluster-level check for this registry is still passing, indicating this is a node-specific issue.
              alert_routing_key: o11y
              team: o11y

  # Node recovers before 30 minutes (should NOT alert)
  - interval: 1m
    input_series:
      - series: 'konflux_up{service="registry-exporter", check="quay.io", source_cluster="c1"}'
        values: '1x31'
      - series: 'registry_exporter_success{tested_registry="quay.io", node="node1", type="pull", source_cluster="c1"}'
        values: '0 0 0 0 0 0 0 0 0 0 0 0 0 0 0 0 0 0 0 0 0 0 0 0 0 0 0 0 0 1 1'
      - series: 'registry_exporter_success{tested_registry="quay.io", node="node1", type="push", source_cluster="c1"}'
        values: '1x31'

    alert_rule_test:
      - eval_time: 30m
        alertname: RegistryExporterNodeFailure
        exp_alerts: []

  # Node has ANY test type failing (not all) - should still alert
  - interval: 1m
    input_series:
      - series: 'konflux_up{service="registry-exporter", check="quay.io", source_cluster="c1"}'
        values: '1x31'
      - series: 'registry_exporter_success{tested_registry="quay.io", node="node1", type="pull", source_cluster="c1"}'
        values: '0x31'
      - series: 'registry_exporter_success{tested_registry="quay.io", node="node1", type="push", source_cluster="c1"}'
        values: '1x31'
      - series: 'registry_exporter_success{tested_registry="quay.io", node="node1", type="metadata", source_cluster="c1"}'
        values: '1x31'
      - series: 'registry_exporter_success{tested_registry="quay.io", node="node1", type="authentication", source_cluster="c1"}'
        values: '1x31'

    alert_rule_test:
      - eval_time: 30m
        alertname: RegistryExporterNodeFailure
        exp_alerts:
          - exp_labels:
              severity: warning
              tested_registry: quay.io
              node: node1
              source_cluster: c1
              slo: "false"
            exp_annotations:
              summary: >-
                Registry exporter running tests on node node1 for registry quay.io in cluster c1 is failing.
              description: >-
                Registry exporter running tests on node node1 for registry quay.io in cluster c1 is failing.
                At least one test type (pull, push, metadata, or authentication) has been failing for 30 minutes.
                The cluster-level check for this registry is still passing, indicating this is a node-specific issue.
              alert_routing_key: o11y
              team: o11y
