rule_files:
  - prometheus.release_service_alerts.yaml

evaluation_interval: 1m

tests:
  - name: ReleaseServiceAlertOnPreProcessingSLOFailure
    interval: 1m
    input_series:
      - series: 'release_pre_processing_duration_seconds_bucket{le="10", job="release", namespace="foo", source_cluster="cluster01"}'
        values: '1+8x59'
      - series: 'release_pre_processing_duration_seconds_count{job="release", namespace="foo", source_cluster="cluster01"}'
        values: '1+9x59'
      - series: 'release_pre_processing_duration_seconds_bucket{le="10", job="release", namespace="foo", source_cluster="cluster01"}'
        values: '1+8x59'
      - series: 'release_pre_processing_duration_seconds_count{job="release", namespace="foo", source_cluster="cluster01"}'
        values: '1+9x59'
    alert_rule_test:
      - eval_time: 1h
        alertname: ReleaseServicePreProcessingDurationSeconds
        exp_alerts:
          - exp_labels:
              severity: high
              slo: "true"
              job: "release"
              source_cluster: "cluster01"
            exp_annotations:
              summary: >-
                90% of Releases must start processing under 10 seconds
              description: >-
                Release service is failing to start processing under 10 seconds for 90% of releases
              alert_team_handle: <!subteam^S03SVBS426R>
              runbook_url: https://gitlab.cee.redhat.com/konflux/docs/sop/-/blob/main/release-service/release-processing-latency.md?ref_type=heads
              team: release

  - name: ReleaseServiceAlertOnValidationSLOFailure
    interval: 1m
    input_series:
      - series: 'release_validation_duration_seconds_bucket{le="5", job="release", namespace="foo", source_cluster="cluster01"}'
        values: '1+8x59'
      - series: 'release_validation_duration_seconds_count{job="release", namespace="foo", source_cluster="cluster01"}'
        values: '1+9x59'
      - series: 'release_validation_duration_seconds_bucket{le="5", job="release", namespace="foo", source_cluster="cluster01"}'
        values: '1+8x59'
      - series: 'release_validation_duration_seconds_count{job="release", namespace="foo", source_cluster="cluster01"}'
        values: '1+9x59'
    alert_rule_test:
      - eval_time: 1h
        alertname: ReleaseServiceValidationDurationSeconds
        exp_alerts:
          - exp_labels:
              severity: critical
              slo: "true"
              job: "release"
              source_cluster: "cluster01"
            exp_annotations:
              summary: >-
                90% of Releases must be validated under 5 seconds
              description: >-
                Release service is failing to run the validations under 5 seconds for 90% of releases
              alert_team_handle: <!subteam^S03SVBS426R>
              runbook_url: https://gitlab.cee.redhat.com/konflux/docs/sop/-/blob/main/release-service/release-validation-latency.md?ref_type=heads
              team: release

  - name: ReleaseServiceControllerAvailabilityAlertOnBelowOf99SLOFailure
    interval: 1m
    input_series:
      - series: 'kube_pod_container_status_ready{namespace="release-service", container="manager", source_cluster="stone-prd-rh01"}'
        values: '1x59'
      - series: 'kube_pod_container_status_ready{namespace="release-service", container="manager", source_cluster="stone-prod-p02"}'
        values: '1x29 0x29'
      - series: 'kube_pod_container_status_ready{namespace="release-service", container="manager", source_cluster="cluster03"}'
        values: '1x59'
    alert_rule_test:
      - eval_time: 1h
        alertname: ReleaseServiceControllerManagerAvailability
        exp_alerts:
          - exp_labels:
              severity: critical
              slo: "true"
              source_cluster: "stone-prod-p02"
            exp_annotations:
              summary: >-
                The Release Service Controller Manager must be available 99% of the time
              description: >-
                The Release Service Controller Manager is failing to be available for 99% of the time
              alert_team_handle: <!subteam^S03SVBS426R>
              runbook_url: https://gitlab.cee.redhat.com/konflux/docs/sop/-/blob/main/release-service/release-controller-availability.md?ref_type=heads
              team: release

  - name: ReleaseServiceControllerAvailabilityNoAlert
    interval: 5m
    input_series:
      - series: 'kube_pod_container_status_ready{namespace="release-service", container="manager", source_cluster="stone-prd-rh01"}'
        values: '1 1x791'
      - series: 'kube_pod_container_status_ready{namespace="release-service", container="manager", source_cluster="stone-prod-p02"}'
        values: '1 1x791'
      - series: 'kube_pod_container_status_ready{namespace="release-service", container="manager", source_cluster="cluster03"}'
        values: '1 1x791'
    alert_rule_test:
      - eval_time: 1h
        alertname: ReleaseServiceControllerManagerAvailability
        exp_alerts: []


  - name: ReleaseServiceControllerManagerOnPodNotReadyStonePrdRh01
    interval: 1m
    input_series:
      - series: 'kube_pod_container_status_ready{namespace="release-service", container="manager", source_cluster="stone-prd-rh01"}'
        values: '1 0x58'
      - series: 'kube_pod_container_status_ready{namespace="release-service", container="manager", source_cluster="stone-prod-p02"}'
        values: '1 1x58'
    alert_rule_test:
      - eval_time: 1h
        alertname: ReleaseServiceControllerManagerPodNotReady
        exp_alerts:
          - exp_labels:
              severity: high
              source_cluster: "stone-prd-rh01"
              namespace: "release-service"
            exp_annotations:
              summary: The Release Service Controller Manager is down in the stone-prd-rh01 for the last 10 minutes
              description: >-
                The Release Service Controller Manager is down in the stone-prd-rh01 for the last 10 minutes
              alert_routing_key: "release-service"
              runbook_url: https://gitlab.cee.redhat.com/konflux/docs/sop/-/blob/main/release-service/release-controller-availability.md?ref_type=heads
              team: release

  - name: ReleaseServiceControllerManagerOnPodNotReadyStoneProdP02
    interval: 1m
    input_series:
      - series: 'kube_pod_container_status_ready{namespace="release-service", container="manager", source_cluster="stone-prd-rh01"}'
        values: '1 1x58'
      - series: 'kube_pod_container_status_ready{namespace="release-service", container="manager", source_cluster="stone-prod-p02"}'
        values: '1 0x58'
    alert_rule_test:
      - eval_time: 1h
        alertname: ReleaseServiceControllerManagerPodNotReady
        exp_alerts:
          - exp_labels:
              severity: high
              source_cluster: "stone-prod-p02"
              namespace: "release-service"
            exp_annotations:
              summary: The Release Service Controller Manager is down in the stone-prod-p02 for the last 10 minutes
              description: >-
                The Release Service Controller Manager is down in the stone-prod-p02 for the last 10 minutes
              alert_routing_key: "release-service"
              runbook_url: https://gitlab.cee.redhat.com/konflux/docs/sop/-/blob/main/release-service/release-controller-availability.md?ref_type=heads
              team: release

  - name: ReleaseServiceControllerManagerOnPodReady
    interval: 1m
    input_series:
      - series: 'kube_pod_container_status_ready{namespace="release-service", container="manager", source_cluster="stone-prd-rh01"}'
        values: '1 1x58'
      - series: 'kube_pod_container_status_ready{namespace="release-service", container="manager", source_cluster="stone-prod-p02"}'
        values: '1 1x58'
    alert_rule_test:
      - eval_time: 1h
        alertname: ReleaseServiceControllerManagerPodNotReady
        exp_alerts: []
