evaluation_interval: 1m

rule_files:
  - 'prometheus.stability_image_repository_provision_failures_alerts.yaml'

tests:
  # Scenario where only one cluster cross the 5 threshold
  # Alert triggered for one cluster
  - interval: 1m
    input_series:
      # Simulating data from Cluster 1 (crosses the 5 threshold)
      - series: 'redhat_appstudio_imagecontroller_image_repository_provision_failure_time_bucket{le="+Inf", source_cluster="cluster01"}'
        values: '0+1x6'  # 7 failed requests

      # Simulating data from Cluster 2 (does not cross the 5 threshold)
      - series: 'redhat_appstudio_imagecontroller_image_repository_provision_failure_time_bucket{le="+Inf", source_cluster="cluster02"}'
        values: '0+1x1'  # 2 failed request

    alert_rule_test:
      - eval_time: 30m
        alertname: StabilityImageRepositoryProvisionFailures
        exp_alerts:
          - exp_labels:
              severity: critical
              slo: "true"
              source_cluster: cluster01
            exp_annotations:
              summary: Stability of image repository provision failures
              description: >
                Provision image repository failures occured for more than 5 requests during the last 30 minutes
                cluster01
              alert_team_handle: <!subteam^S0500D143S8>
              team: build
              runbook_url: https://gitlab.cee.redhat.com/konflux/docs/sop/-/blob/main/image-controller/stability_image_repository_provision_failures.md

  # Scenario where both clusters cross the 5 threshold
  # Alert triggered for both clusters
  - interval: 1m
    input_series:
      # Simulating data from Cluster 1 (crosses the 5 threshold)
      - series: 'redhat_appstudio_imagecontroller_image_repository_provision_failure_time_bucket{le="+Inf", source_cluster="cluster01"}'
        values: '0+1x6'  # 7 failed requests

      # Simulating data from Cluster 2 (also crosses the 5 threshold)
      - series: 'redhat_appstudio_imagecontroller_image_repository_provision_failure_time_bucket{le="+Inf", source_cluster="cluster02"}'
        values: '0+1x6'  # 7 failed requests

    alert_rule_test:
      - eval_time: 30m
        alertname: StabilityImageRepositoryProvisionFailures
        exp_alerts:
          - exp_labels:
              severity: critical
              slo: "true"
              source_cluster: cluster01
            exp_annotations:
              summary: Stability of image repository provision failures
              description: >
                Provision image repository failures occured for more than 5 requests during the last 30 minutes
                cluster01
              alert_team_handle: <!subteam^S0500D143S8>
              team: build
              runbook_url: https://gitlab.cee.redhat.com/konflux/docs/sop/-/blob/main/image-controller/stability_image_repository_provision_failures.md
          - exp_labels:
              severity: critical
              slo: "true"
              source_cluster: cluster02
            exp_annotations:
              summary: Stability of image repository provision failures
              description: >
                Provision image repository failures occured for more than 5 requests during the last 30 minutes
                cluster02
              alert_team_handle: <!subteam^S0500D143S8>
              team: build
              runbook_url: https://gitlab.cee.redhat.com/konflux/docs/sop/-/blob/main/image-controller/stability_image_repository_provision_failures.md

  # Scenario where neither cluster crosses the 5 threshold
  # Alert not triggered
  - interval: 1m
    input_series:
      # Simulating data from Cluster 1 (does not cross the 5 threshold)
      - series: 'redhat_appstudio_imagecontroller_image_repository_provision_failure_time_bucket{le="+Inf", source_cluster="cluster01"}'
        values: '0+1x1'  # 2 failed request

      # Simulating data from Cluster 2 (does not cross the 5 threshold)
      - series: 'redhat_appstudio_imagecontroller_image_repository_provision_failure_time_bucket{le="+Inf", source_cluster="cluster02"}'
        values: '0+1x1'  # 2 failed request

    alert_rule_test:
      - eval_time: 30m
        alertname: StabilityImageRepositoryProvisionFailures
        exp_alerts: []  # No alerts are expected in this scenario
