evaluation_interval: 1m

rule_files:
  - prometheus.argocd_alerts.yaml

tests:
  # DegradedArgocdApp alert
  - interval: 1m
    input_series:

      # no app is degraded - shouldn't trigger alerts
      - series: 'argocd_app_info{health_status="Degraded", name="not-degraded-app", namespace="argocd-staging", dest_namespace="test_dest_namespace", dest_server="https://api.foo.openshiftapps.com:6443", source_environment="staging-cluster"}'
        values: _x9

      # an app is constantly degraded - should trigger alerts
      - series: 'argocd_app_info{health_status="Degraded", name="degraded-app", namespace="argocd-staging", dest_namespace="test_dest_namespace", dest_server="https://api.foo.openshiftapps.com:6443", source_environment="staging-cluster"}'
        values: 1x9

      # an app with some 1m degraded status - shouldn't trigger alerts
      - series: 'argocd_app_info{health_status="Degraded", name="degraded-1m-app", namespace="argocd", dest_namespace="test_dest_namespace", dest_server="https://api.foo.openshiftapps.com:6443", source_environment="production-cluster"}'
        values: 1 _ _ _ 1 _ _ _ _ _

      # an app with a 4m degraded status - shouldn't trigger alerts
      - series: 'argocd_app_info{health_status="Degraded", name="degraded-4m-app", namespace="argocd", dest_namespace="test_dest_namespace", dest_server="https://api.foo.openshiftapps.com:6443", source_environment="production-cluster"}'
        values: _ 1 1 1 1 _ _ _ _ _

      # an app is flapping every 1m - should trigger alerts triggers
      - series: 'argocd_app_info{health_status="Degraded", name="flapping-1m-app", namespace="argocd", dest_namespace="test_dest_namespace", dest_server="https://api.foo.openshiftapps.com:6443", source_environment="production-cluster"}'
        values: 1 _ 1 _ 1 _ 1 _ 1 _

      # an app is flapping every 2m - should trigger alerts triggers
      - series: 'argocd_app_info{health_status="Degraded", name="flapping-2m-app", namespace="argocd", dest_namespace="test_dest_namespace", dest_server="https://api.foo.openshiftapps.com:6443", source_environment="production-cluster"}'
        values: 1 _ _ 1 _ _ 1 _ _ 1

      # an app is flapping every 2m but on arbitrary namespace - should not trigger
      - series: 'argocd_app_info{health_status="Degraded", name="flapping-2m-app", namespace="baz", dest_namespace="test_dest_namespace", dest_server="https://api.foo.openshiftapps.com:6443", source_environment="production-cluster"}'
        values: 1 _ _ 1 _ _ 1 _ _ 1

      # an app is in another health status which is not 'Degraded' - the rule shouldn't alert.
      - series: 'argocd_app_info{health_status="Progressing", name="progressing-app", namespace="argocd-staging", dest_namespace="test_dest_namespace", dest_server="https://api.foo.openshiftapps.com:6443", source_environment="staging-cluster"}'
        values: 1x9


    alert_rule_test:
      - eval_time: 5m
        alertname: DegradedArgocdApp
        exp_alerts:
          - exp_labels:
              severity: warning
              dest_namespace: test_dest_namespace
              dest_server: https://api.foo.openshiftapps.com:6443
              namespace: argocd-staging
              health_status: Degraded
              name: degraded-app
              source_environment: staging-cluster
              team: rhtap
              control_plane: "true"
            exp_annotations:
              summary: Argocd App 'degraded-app' is degraded
              description: |
                Environment: staging-cluster
                Application: degraded-app
                Cluster: https://api.foo.openshiftapps.com:6443
                Health status Degraded.
              alert_routing_key: degraded-app
              runbook_url: https://gitlab.cee.redhat.com/rhtap/docs/sop/-/blob/main/o11y/alert-rule-degradedArgocdApp.md

          - exp_labels:
              severity: warning
              dest_namespace: test_dest_namespace
              dest_server: https://api.foo.openshiftapps.com:6443
              namespace: argocd
              health_status: Degraded
              name: flapping-1m-app
              source_environment: production-cluster
              team: rhtap
              control_plane: "true"
            exp_annotations:
              summary: Argocd App 'flapping-1m-app' is degraded
              description: |
                Environment: production-cluster
                Application: flapping-1m-app
                Cluster: https://api.foo.openshiftapps.com:6443
                Health status Degraded.
              alert_routing_key: flapping-1m-app
              runbook_url: https://gitlab.cee.redhat.com/rhtap/docs/sop/-/blob/main/o11y/alert-rule-degradedArgocdApp.md

          - exp_labels:
              severity: warning
              dest_namespace: test_dest_namespace
              dest_server: https://api.foo.openshiftapps.com:6443
              namespace: argocd
              health_status: Degraded
              name: flapping-2m-app
              source_environment: production-cluster
              team: rhtap
              control_plane: "true"
            exp_annotations:
              summary: Argocd App 'flapping-2m-app' is degraded
              description: |
                Environment: production-cluster
                Application: flapping-2m-app
                Cluster: https://api.foo.openshiftapps.com:6443
                Health status Degraded.
              alert_routing_key: flapping-2m-app
              runbook_url: https://gitlab.cee.redhat.com/rhtap/docs/sop/-/blob/main/o11y/alert-rule-degradedArgocdApp.md

  # ProgressingArgocdApp alert
  - interval: 1m
    input_series:
      # App was progessing and then healthy for the last 9 min, should not raise an alert
      - series: 'argocd_app_info{health_status="Progressing", name="Progressing-Healthy-9min-app", namespace="argocd-staging", dest_server="https://api.foo.openshiftapps.com:6443", source_environment="staging-cluster"}'
        values: 1x10 _x8
      - series: 'argocd_app_info{health_status="Healthy", name="Progressing-Healthy-9min-app", namespace="argocd-staging", dest_server="https://api.foo.openshiftapps.com:6443", source_environment="staging-cluster"}'
        values: _x10 1x8
    alert_rule_test:
      - eval_time: 20m
        alertname: ProgressingArgocdApp
        exp_alerts: []

  - interval: 1m
    input_series:
      # App is always progressing, should raise an alert
      - series: 'argocd_app_info{health_status="Progressing", name="progressing-only-app", namespace="argocd-staging", dest_namespace="test_dest_namespace", dest_server="https://api.foo.openshiftapps.com:6443", source_environment="staging-cluster"}'
        values: 1x19
    alert_rule_test:
      - eval_time: 20m
        alertname: ProgressingArgocdApp
        exp_alerts:
          - exp_labels:
              severity: warning
              dest_namespace: test_dest_namespace
              dest_server: https://api.foo.openshiftapps.com:6443
              namespace: argocd-staging
              health_status: Progressing
              name: progressing-only-app
              source_environment: staging-cluster
              team: rhtap
              control_plane: "true"
            exp_annotations:
              summary: Argocd App 'progressing-only-app' is progressing for too long
              description: |-
                Environment: staging-cluster
                Application: progressing-only-app
                Cluster: https://api.foo.openshiftapps.com:6443
                App progressing for too long.
              alert_routing_key: progressing-only-app
              runbook_url: https://gitlab.cee.redhat.com/rhtap/docs/sop/-/blob/main/o11y/alert-rule-ProgressingArgocdApp.md

  - interval: 1m
    input_series:
      # App is always progressing, but on unrelated namespace. should NOT raise an alert
      - series: 'argocd_app_info{health_status="Progressing", name="progressing-only-app", namespace="not-us", dest_server="https://api.foo.openshiftapps.com:6443", source_environment="staging-cluster"}'
        values: 1x19
    alert_rule_test:
      - eval_time: 20m
        alertname: ProgressingArgocdApp
        exp_alerts: []

  - interval: 1m
    input_series:
      # Apps are always healthy, should not raise an alert
      - series: 'argocd_app_info{health_status="Healthy", name="Healthy-only-app", namespace="argocd", dest_server="https://api.foo.openshiftapps.com:6443", source_environment="production-cluster"}'
        values: 1x19
      - series: 'argocd_app_info{health_status="Progressing", name="Healthy-only-app", namespace="argocd", dest_server="https://api.foo.openshiftapps.com:6443", source_environment="production-cluster"}'
        values: _x19
      - series: 'argocd_app_info{health_status="Healthy", name="another-Healthy-only-app", namespace="argocd", dest_server="https://api.foo.openshiftapps.com:6443", source_environment="production-cluster"}'
        values: 1x19
    alert_rule_test:
      - eval_time: 20m
        alertname: ProgressingArgocdApp
        exp_alerts: []

  - interval: 1m
    input_series:
      # App was progessing and then healthy for the last 4 min, should not raise an alert
      - series: 'argocd_app_info{health_status="Progressing", name="progressing-healthy-4min-app", namespace="argocd", dest_server="https://api.foo.openshiftapps.com:6443", source_environment="production-cluster"}'
        values: 1x15 _x3
      - series: 'argocd_app_info{health_status="Healthy", name="progressing-healthy-4min-app", namespace="argocd", dest_server="https://api.foo.openshiftapps.com:6443", source_environment="production-cluster"}'
        values: _x15 1x3
    alert_rule_test:
      - eval_time: 20m
        alertname: ProgressingArgocdApp
        exp_alerts: []

  - interval: 1m
    input_series:
      # the app is switching health states (flapping), should not raise an alert
      - series: 'argocd_app_info{health_status="Progressing", name="flapping-should-not-alert-app", namespace="argocd", dest_server="https://api.foo.openshiftapps.com:6443", source_environment="production-cluster"}'
        values: '_ 1 _ 1 1 _ 1 1 1 _ _ _ _ 1 _ 1 1 1 1 _'
      - series: 'argocd_app_info{health_status="Healthy", name="flapping-should-not-alert-app", namespace="argocd", dest_server="https://api.foo.openshiftapps.com:6443", source_environment="production-cluster"}'
        values: '1 _ 1 _ _ 1 _ _ _ 1 1 1 1 _ 1 _ _ _ _ 1'
    alert_rule_test:
      - eval_time: 20m
        alertname: ProgressingArgocdApp
        exp_alerts: []

  - interval: 1m
    input_series:
      # App 1: Progressing all of the time, should raise an alert
      - series: 'argocd_app_info{health_status="Progressing", name="Progressing_first-app", namespace="argocd", dest_namespace="test_dest_namespace", dest_server="https://api.foo.openshiftapps.com:6443", source_environment="production-cluster"}'
        values: 1x19
      - series: 'argocd_app_info{health_status="Healthy", name="Progressing_first-app", namespace="argocd", dest_namespace="test_dest_namespace", dest_server="https://api.foo.openshiftapps.com:6443", source_environment="production-cluster"}'
        values: _x19
      # App 2: Progressing for 17 minutes and then healthy for 3 minutes, should not raise an alert
      - series: 'argocd_app_info{health_status="Progressing", name="Progressing_second-app", namespace="argocd", dest_namespace="test_dest_namespace", dest_server="https://api.foo.openshiftapps.com:6443", source_environment="production-cluster"}'
        values: 1x16 _ _ _
      - series: 'argocd_app_info{health_status="Healthy", name="Progressing_second-app", namespace="argocd", dest_namespace="test_dest_namespace", dest_server="https://api.foo.openshiftapps.com:6443", source_environment="production-cluster"}'
        values: _x16 1 1 1
    alert_rule_test:
      - eval_time: 20m
        alertname: ProgressingArgocdApp
        exp_alerts:
          - exp_labels:
              severity: warning
              dest_namespace: test_dest_namespace
              dest_server: https://api.foo.openshiftapps.com:6443
              namespace: argocd
              health_status: Progressing
              name: Progressing_first-app
              source_environment: production-cluster
              team: rhtap
              control_plane: "true"
            exp_annotations:
              summary: Argocd App 'Progressing_first-app' is progressing for too long
              description: |-
                Environment: production-cluster
                Application: Progressing_first-app
                Cluster: https://api.foo.openshiftapps.com:6443
                App progressing for too long.
              alert_routing_key: Progressing_first-app
              runbook_url: https://gitlab.cee.redhat.com/rhtap/docs/sop/-/blob/main/o11y/alert-rule-ProgressingArgocdApp.md
